HiExM: high-throughput expansion microscopy enables scalable super-resolution imaging

Author response:

The following is the authors’ response to the previous reviews.

Public Reviews:

Reviewer #1 (Public review):

Summary

In this manuscript, Day et al. present a high-throughput version of expansion microscopy to increase the throughput of this well-established super-resolution imaging technique. Through technical innovations in liquid handling with custom-fabricated tools and modifications to how the expandable hydrogels are polymerized, the authors show robust ~4-fold expansion of cultured cells in 96-well plates. They go on to show that HiExM can be used for applications such as drug screens by testing the effect of doxorubicin on human cardiomyocytes. Interestingly, the effects of this drug on changing DNA organization were only detectable by ExM, demonstrating the utility of HiExM for such studies.

Overall, this is a very well-written manuscript presenting an important technical advance that overcomes a major limitation of ExM - throughput. As a method, HiExM appears extremely useful and the data generally support the conclusions.

Strengths

Hi-ExM overcomes a major limitation of ExM by increasing the throughput and reducing the need for manual handling of gels. The authors do an excellent job of explaining each variation introduced to HiExM to make this work and thoroughly characterize the impressive expansion isotropy. The dox experiments are generally well-controlled and the comparison to an alternative stressor (H2O2) significantly strengthens the conclusions.

Weaknesses

(1) It is still unclear to me whether or not cells that do not expand remain in the well given the response to point 1. The authors say the cells are digested and washed away but then say that there is a remaining signal from the unexpanded DNA in some cases. I believe this is still a concern that potential users of the protocol should be aware of.

Although ProteinaseK digestion removes most of the unexpanded cells, DNA can sometimes persist. As such, we occasionally observe Hoechst signal underneath cells. The residual DNA is easily differentiated from nuclear Hoechst signal and does not confound interpretation of results. We have added a new supplementary figure that further clarifies this point.

(2) Regarding the response to point 9, I think this information should be included in the manuscript, possibly in the methods. It is important for others to have a sense of how long imaging may take if they were to adopt this method.

We have added detailed information to the methods section to address this point as shown below. In general, we image HiExM samples on the Opera Phenix at 63x with the following parameters: 100% laser power for all channels; 200 ms exposure for Hoechst, 500-1000+ ms exposure for immunostained channels depending on the strength of the stain and the laser; 60 optical sections with 1 micron spacing; and 4-20 fields of view per well depending on the cell density and sample size requirements. Therefore, imaging one full 96-well plate (60 wells total as we avoid the outer wells) takes anywhere from 3 hr to 64 hr depending on the combination of parameters used.

Reviewer #2 (Public review):

Summary:

In the present work, the authors present an engineering solution to sample preparation in 96-well plates for high-throughput super resolution microscopy via Expansion Microscopy. This is not a trivial problem, as the well cannot be filled with the gel, which would prohibit expansion of the gel. They thus engineered a device that can spot a small droplet of hydrogel solution and keep it in place as it polymerises. It occupies only a small portion space at the center of each well, the gel can expand into all directions and imaging and staining can proceed by liquid handling robots and an automated microscope.

Strengths:

In contrast to Reference 8, the authors system is compatible with standard 96 well imaging plates for high-throughput automated microscopy and automated liquid handling for most parts of the protocol. They thus provide a clear path towards high throughput exM and high throughout super resolution microscopy, which is a timely and important goal.

Addition upon revision:

The authors addressed this reviewer's suggestions.

Reviewer #3 (Public review):

Summary:

Day et al. introduced high-throughput expansion microscopy (HiExM), a method facilitating the simultaneous adaptation of expansion microscopy for cells cultured in a 96-well plate format. The distinctive features of this method include: 1) the use of a specialized device for delivering a minimal amount (~230 nL) of gel solution to each well of a conventional 96-well plate, and 2) the application of the photochemical initiator, Irgacure 2959, to successfully form and expand toroidal gel within each well.

Addition upon revision:

Overall, the authors have adequately addressed most of the concerns raised. There are a few minor issues that require attention.

Minor comments:

Figure S10: There appears to be a discrepancy in the panel labeling. The current labels are EH, but it is unclear whether panels A-D exist. Also, this reviewer thought that panels G and H would benefit from statistical testing to strengthen the conclusions. As a general rule for scientific graph presentation, the y-axis of all graphs should start at zero unless there is a compelling reason not to do so.

We have revised Figure S10 to address your comments.

Changes from the earlier version to the revised version

Abstract

Abstract

Expansion microscopy (ExM) enables nanoscale imaging using a standard confocal microscope through the physical, isotropic expansion of fixed immunolabeled specimens. ExM is widely employed to image proteins, nucleic acids, and lipid membranes in single cells; however, current methods limit the number of samples that can be processed simultaneously. We developed High-throughput Expansion Microscopy (HiExM), a robust platform that enables expansion microscopy of cells cultured in a standard 96-well plate. Our method enables ∼4.2x expansion of cells within individual wells, across multiple wells, and between plates. We also demonstrate that HiExM can be combined with high-throughput confocal imaging platforms to greatly improve the ease and scalability of image acquisition. As an example, we analyzed the effects of doxorubicin, a known cardiotoxic agent, on human cardiomyocytes (CMs) as measured by Hoechst signal across the nucleus. We show a dose dependent effect on nuclear DNA that is not observed in unexpanded CMs, suggesting that HiExM improves the detection of cellular phenotypes in response to drug treatment. Our method broadens the application of ExM as a tool for scalable super-resolution imaging in biological research applications.

Main Text

The ability to resolve subcellular features can greatly enhance biological discovery, identification of disease targets, and the development of targeted therapeutics. Super-resolution microscopy methods such as structured illumination microscopy (SIM), stochastic optical reconstruction microscopy (STORM) and stimulated emission depletion microscopy (STED) enable nanoscale imaging of specimens.1–3 However, these approaches require specialized expertise, costly reagents and microscopes. Expansion microscopy (ExM) is an inexpensive and accessible method that similarly resolves fluorescently labeled structures at subcellular resolution by isotropically enlarging specimens within a swellable polyelectrolyte hydrogel.4 Images of expanded cells show a higher effective resolution due to an increase in the distance between the fluorescent molecules. Super-resolution imaging and ExM are critical tools that are widely used in biological research; however, sample processing and imaging require time consuming manual manipulation, limiting throughput. Although there has been success in increasing the throughput of certain super-resolution microscopy techniques, these methods require highly specialized equipment and expertise which limits the accessibility of these tools.5–8 To circumvent these limitations, we developed High-Throughput Expansion Microscopy (HiExM) that combines parallel sample processing and automated high-content confocal imaging in a standard 96-well cell culture plate.

Adapting standard ExM to a 96-well plate requires the reproducible delivery of a small volume of gel solution to each well (<1 µL compared to ∼200 µL/sample for slide-based ExM). This small volume is necessary to allow complete gel expansion within the size constraints of wells in a 96-well plate. To achieve this goal, we engineered a device that retains small liquid droplets when dipped into a reservoir of monomeric expansion gel solution. These droplets can then be deposited into the centers of wells across the cell culture plate (Figs. 1a,b). The retention of the gel solution is mediated by a series of reticulated grooves at the end of a cylindrical post which allows for the formation of a pendant droplet (Fig. 1c). We estimate that each post retains ∼230 nL of gel solution as measured by the mass of liquid accumulated on all post tips of a single device after dipping (Supplementary Methods).

HiExM enables gel formation and expansion in a 96-well cell culture plate

a. Schematic representation of HiExM devices showing the key features highlighted in color. b. Example devices used in 96-well cell culture plates. c. Brightfield image of the conical post-tip shows the pattern of grooves that mediate fluid retention. d. Fluid retention at the conical post-tip of the device. Silhouettes taken by an optical comparator of the profile of a single post suspended above a surface (left) and in contact with a surface (right) show a fluid droplet interacting with the device. Upon device insertion, the gel solution fills the space under the conical post tip, forming the toroid gel. e-h. Schematic of HiExM gel deposition and expansion workflow. e. The device is immersed in a shallow reservoir of gel solution. f. Upon removal, the tip of each device post retains a small volume of gel solution. g. Gel solution is deposited by the device into the well centers of the cell culture plate. Brightfield image (right) shows gel geometry and size prior to expansion. Note that gels deposited in HiExM cover ∼1.1 mm2 of the well surface to accommodate the expanded gel, and do not include cells outside the gel footprint. h. Polymerization and expansion are performed with the device in place. Brightfield image (right) shows gel geometry and size after expansion.

The small volume of gel solution delivered to each well by the device forms a toroidal droplet where the inner surface is molded by the conical surface of the post-tip and the outer surface is constrained by surface tension (Figs. 1c,d). The post-tip dimensions allow delivery of a volume of gel solution that expands relative to the diameter of the well (Fig. 1d). The toroidal gel enables isotropic gel expansion within the well constraints while allowing for easy device removal because the inner surface of the gel expands outward away from the conical post-tip, effectively detaching the gel from the device. To achieve this toroidal geometry, our device design incorporates three key features to ensure each post of the device contacts the midpoint at the bottom of each well (Fig. S1). First, two sets of three cantilevers (‘pressure struts’) are located at each end of the device and flank two rows of 6 posts (12 posts per device). These pressure struts apply outward force to the inside of the outer four wells ensuring that the device remains in contact with the well surface throughout the expansion process. Second, the spine of the device incorporates several notches extending through the long side of the device allowing it to bend to accommodate plate inconsistencies between the depths of individual wells. Finally, the post lengths are offset from the center four posts to allow for sequential deformation that ensures each post contacts the well surface. Together, these design features allow for reproducible deposition and expansion of gels in a 96-well plate.

To test our device for high-throughput gel expansion, we first performed ProExM using ammonium persulfate (APS) and tetramethylethylenediamine (TEMED) to initiate gel polymerization (Fig. S2a).9 Since APS rapidly accelerates polymerization of polyacrylamide hydrogels, TEMED-containing expansion gel solution is first delivered by the device and removed leaving a ≤1µl droplet in each well of the plate on ice. Another device is then used to deliver APS-containing expansion gel solution at the same volume and the polymerization reaction is initiated by placing the plate on a 50°C surface. Given the large air-liquid-interface of the gel, droplet delivery and polymerization are performed in a nitrogen filled glove bag to minimize oxygen inhibition of the reaction. Although we observed gel formation across wells, polymerization and expansion were inconsistent, likely due to inconsistent radical formation and rapid evaporation of the small volume relative to the rate of polymerization (Movie S1). To overcome this limitation, we tested photochemical initiators including Lithium phenyl(2,4,6-trimethylbenzoyl)phosphinate (LAP), diphenyl(2,4,6-trimethylbenzoyl)phosphine oxide (TPO), and 1-[4-(2-Hydroxyethoxy)-phenyl]-2-hydroxy-2-methyl-1-propane-1-one (Irgacure 2959) which control polymerization of polyacrylamide gels when exposed to UV light.10,11 LAP has shown success in slide-based ExM using PEG-based hydrogels.12,13 We found that Irgacure 2959 allowed reproducible gel formation and expansion when compared to either APS/TEMED or other photoinitiators (Movie S2).

Earlier version: We next tested the ability of the HiExM platform to achieve nanoscale image resolution of biological specimens. Briefly, A549 cells were cultured, fixed, permeabilized and immunostained with alpha-tubulin antibodies to visualize microtubules as a reference structure before and after expansion. Stained samples were then incubated with Acryloyl-X (AcX) overnight at 4°C to allow anchoring of native proteins to the polymer matrix during gel formation. Device posts were immersed in the expansion gel solution containing 0.1% Irgacure 2959. Devices were inserted into the well plate followed by exposure to UV light (365nm) in an anoxic environment at room temperature to initiate the polymerization reaction. Cells embedded in the resulting gels were digested with Proteinase K then expanded in deionized water overnight (Fig. S2b, Fig. S3). Visual inspection using an epifluorescence microscope showed both robust gel formation and fluorescence signal retention across wells. Although we observed bleaching of AlexaFluor® dyes in our Irgacure 2959 photopolymerized gels, Cyanine-based Fluorescent (CF®) dyes are robust against bleaching in HiExM and result in reproducible signal retention (Fig. S4).14 We also titrated AcX and ProteinaseK to optimize signal retention in HiExM and found the 50 µg/mL AcX and 1U/mL ProteinaseK was most consistent for A549 cells (Fig. S5).

Revised version: We next tested the ability of the HiExM platform to achieve nanoscale image resolution of biological specimens. Briefly, A549 cells were cultured, fixed, permeabilized and immunostained with alpha-tubulin antibodies to visualize microtubules as a reference structure before and after expansion. Stained samples were then incubated with Acryloyl-X (AcX) overnight at 4°C to allow anchoring of native proteins to the polymer matrix during gel formation. Device posts were immersed in the expansion gel solution containing 0.1% Irgacure 2959. Devices were inserted into the well plate followed by exposure to UV light (365nm) in an anoxic environment at room temperature to initiate the polymerization reaction. Cells embedded in the resulting gels were digested with Proteinase K then expanded in deionized water overnight (Fig. S2b, Fig. S3). Visual inspection using an epifluorescence microscope showed both robust gel formation and fluorescence signal retention across wells. Although we observed bleaching of AlexaFluorⓇ dyes in our Irgacure 2959 photopolymerized gels, Cyanine-based Fluorescent (CFⓇ) dyes are robust against bleaching in HiExM and result in reproducible signal retention (Fig. S4).14 We also titrated AcX and ProteinaseK to optimize signal retention in HiExM and found the 50 µg/mL AcX and 1U/mL ProteinaseK was most consistent for A549 cells (Fig. S5). This optimization step is likely critical for different cell types. In some cases, we observed residual Hoechst signal underneath expanded cells due to digested cells that were incompletely removed during wash steps during the expansion process (Fig. S6). This signal does not impact interpretation of results.

Earlier version: A major workflow bottleneck associated with super-resolution imaging, including current ExM protocols, is the significant time and expertise required for image acquisition. For example, the vertical component of expansion increases the depth of the sample, necessitating more optical z-slices to image the full depth of the sample. Moreover, expansion effectively decreases the field-of-view due to increasing sample size, necessitating acquisition of more images to analyze the same number of cells at lower resolution. To enable rapid 3D image capture of expanded cells in 96-well plates, we used a high content confocal microscope (Opera Phenix system, Perkin Elmer) (Fig. 2a). This approach allowed autonomous capture of images across the entire plate at a resolution of ∼115 nm compared to ∼463 nm for unexpanded samples using the same microscope, objective (63X, 1.15NA water immersion), and imaging wavelength (Fig. S6). Image resolution was calculated using a decorrelation analysis.15 We used the PreciScan plug-in in the Harmony analysis software to first accurately target regions of interest based on Hoechst staining to reduce acquisition times at higher magnification. Briefly, we imaged plates at 5x across entire wells and stored the coordinates of nuclei for each well. These coordinates were then used to direct the 63x magnification objective to the relevant fields of view.

Revised version: A major workflow bottleneck associated with super-resolution imaging, including current ExM protocols, is the significant time and expertise required for image acquisition. For example, the vertical component of expansion increases the depth of the sample, necessitating more optical z-slices to image the full depth of the sample. Moreover, expansion effectively decreases the field-of-view due to increasing sample size, necessitating acquisition of more images to analyze the same number of cells at lower resolution. To enable rapid 3D image capture of expanded cells in 96-well plates, we used a high content confocal microscope (Opera Phenix system, Perkin Elmer) (Fig. 2a). This approach allowed autonomous capture of images across the entire plate at a resolution of ∼115 nm compared to ∼463 nm for unexpanded samples using the same microscope, objective (63X, 1.15NA water immersion), and imaging wavelength (Fig. S7). Image resolution was calculated using a decorrelation analysis.15 We used the PreciScan plug-in in the Harmony analysis software to first accurately target regions of interest based on Hoechst staining to reduce acquisition times at higher magnification. Briefly, we imaged plates at 5x across entire wells and stored the coordinates of nuclei for each well. These coordinates were then used to direct the 63x magnification objective to the relevant fields of view.

HiExM is compatible with nanoscale expansion and automated image acquisition of human cells with minimal distortion

a. Expanded gels in a 96 well plate imaged at 5x with an inset showing an individual gel shown on the right. b. Imaging fields shown for pre-expansion and post-expansion samples. Prior to expansion, 61 fields were imaged at 63x around the center of the gel. After expansion, 12 fields were imaged at 20x magnification. Max projections for each field were then stitched together and manually inspected to identify fields in the pre-expansion and post-expansion image sets containing matching cells. Arrows show an example of a field representing the same cells before and after expansion. c. Representative registered fluorescence image of microtubules before (left) and after (right) expansion showing the increase in resolution conferred by expansion. In this case, both pre-expansion and post-expansion images were taken at 63x magnification. d. Representative error curve calculated using non-rigid registration of 43 independent fields of view. The shaded region denotes one standard error of the mean. e. Error is not dependent on the location of a given field of view within the gel. Each data point represents the average percent error (up to 40 μm measurement length post-expansion) for a given field of view as determined by an NRR analysis. Seven to ten fields of view were analyzed across six wells. The distance from the center of each field of view to the center of the well was measured manually in FIJI using the stitched pre-expansion image set. f. Box plots showing the distribution of average error among gels across three plates (average = 3.63% +/− 1.39). Average error percentages for each well plate: 3.73 +/− 0.879, 3.89 +/− 1.65, and 3.43 +/− 1.44. g. Expansion factor is not dependent on the location of a given field of view within the gel. Images taken before and after expansion of the same cell or group of cells were used to measure the ratio of length scales in the two images. h. Box plots showing the distribution of expansion factors of 58 gels across three plates. Expansion factors ranged from 3.5-5.1x across gels with an average of 4.16 +/− 0.394. Average expansion factors for each well plate showed tight distribution with some variability across plates: 3.78 +/− 0.171, 4.26 +/− 0.171, and 4.53 +/− 0.321. Comparison of the three well plates was performed using one-way ANOVA and post-hoc Tukey HSD test (*p-value<<0.05).

Earlier version: We next benchmarked the performance of our platform relative to slide-based proExM.9 Image distortion due to expansion was measured using a non-rigid registration (NRR) algorithm.4 Prior to the application of HiExM, images of microtubule stained A549 cells were acquired using the Opera Phenix from 61 fields around the center of each well (Fig. 2a). Cells were then expanded as above and 12 fields were imaged per well. Images from pre- and post-expansion were stitched together and the images for each well were then compared side-by-side to identify matched cells. Although we observed that gels did not maintain a fixed position pre- and post-expansion, close inspection of the images allowed us to align their orientations. Distortion was calculated using the root mean squared error (RMSE) of a 2-dimensional deformation vector field by comparing microtubule morphology in the same cell before and after expansion in six to ten fields in each well across six arbitrarily selected wells (Figs. 2c,d). The average percent error (up to 40 μm measurement length post-expansion) was plotted against the distance from the center of the gel (Fig. 2e). Our analysis found no correlation (R2 = 0.003) between expansion error and distance from the gel center. In fact, we observed only a small number of outliers (>3 s.d.) in any given well (<2 instances per gel across 6 gels). These outliers are not expected to affect biological interpretation. The outliers, which are closer to the center of the gel, likely occur when the device post is not in full contact with the well surface, resulting in a higher volume of gel at the center that resists gel swelling under the device (Fig. S7). Expansion error was also highly consistent when comparing wells across multiple plates (average = 3.68%), comparable to slide-based ExM (Fig. 2f).4,9 We observed only a small number of image fields (7 of 58) with average percent error above 6% (maximum 6.78%). Moreover, calculated error showed that photoinitiation resulted in markedly lower spatial errors compared to APS and TEMED polymerization chemistry (Fig. S8). Analyzing cell seeding densities also showed that cell confluency did not affect image distortion as measured by NRR (Fig. S9).

Revised version: We next benchmarked the performance of our platform relative to slide-based proExM.9 Image distortion due to expansion was measured using a non-rigid registration (NRR) algorithm.4 Prior to the application of HiExM, images of microtubule stained A549 cells were acquired using the Opera Phenix from 61 fields around the center of each well (Fig. 2a). Cells were then expanded as above and 12 fields were imaged per well. Images from pre- and post-expansion were stitched together and the images for each well were then compared side-by-side to identify matched cells. Although we observed that gels did not maintain a fixed position pre- and post-expansion, close inspection of the images allowed us to align their orientations. Distortion was calculated using the root mean squared error (RMSE) of a 2-dimensional deformation vector field by comparing microtubule morphology in the same cell before and after expansion in six to ten fields in each well across six arbitrarily selected wells (Figs. 2c,d). The average percent error (up to 40 μm measurement length post-expansion) was plotted against the distance from the center of the gel (Fig. 2e). Our analysis found no correlation (R2 = 0.003) between expansion error and distance from the gel center. In fact, we observed only a small number of outliers (>3 s.d.) in any given well (<2 instances per gel across 6 gels). These outliers are not expected to affect biological interpretation. The outliers, which are closer to the center of the gel, likely occur when the device post is not in full contact with the well surface, resulting in a higher volume of gel at the center that resists gel swelling under the device (Fig. S8). Expansion error was also highly consistent when comparing wells across multiple plates (average = 3.68%), comparable to slide-based ExM (Fig. 2f).4,9 We observed only a small number of image fields (7 of 58) with average percent error above 6% (maximum 6.78%). Moreover, calculated error showed that photoinitiation resulted in markedly lower spatial errors compared to APS and TEMED polymerization chemistry (Fig. S9). Analyzing cell seeding densities also showed that cell confluency did not affect image distortion as measured by NRR (Fig. S10).

Earlier version: We next measured expansion factors within individual gels, between gels, and across plates by manually measuring the distance between matched pairs of cell landmarks in FIJI before and after expansion. We found that the expansion factor was highly reproducible across different regions of the gel confirming robust expansion using our method (Fig. 2g). The average expansion factor of gels sampled across three plates was 4.16 +/- 0.394 (Fig. 2h), consistent with proExM.9 We next compared the distributions of expansion factors across 3 well plates by ANOVA and found statistically significant differences between expansion factor distributions.

Revised version: We next measured expansion factors within individual gels, between gels, and across plates by manually measuring the distance between matched pairs of cell landmarks in FIJI before and after expansion. We found that the expansion factor was highly reproducible across different regions of the gel confirming robust expansion using our method (Fig. 2g). The average expansion factor of gels sampled across three plates was 4.16 +/− 0.394 (Fig. 2h), consistent with proExM.9 We next compared the distributions of expansion factors across 3 well plates by ANOVA and found statistically significant differences between expansion factor distributions. Thus, while uniform within gels, the expansion factor is somewhat variable between gels and plates. We attribute these differences primarily to the small size of the gels, making them vulnerable to the effects of evaporation. This variability should be taken into consideration for studies where absolute length measurements between plates are important for biological interpretation. Collectively, these data show that the HiExM method results in reproducible nanoscale and isotropic expansion of cells within gels, across wells, and plates.

Removed in the revised version: Thus, while uniform within gels, the expansion factor is somewhat variable between gels and plates. We attribute these differences primarily to the small size of the gels, making them vulnerable to the effects of evaporation. This variability should be taken into consideration for studies where absolute length measurements between plates are important for biological interpretation. Collectively, these data show that the HiExM method results in reproducible nanoscale and isotropic expansion of cells within gels, across wells, and plates.

Earlier version: The ability to image cells at increased resolution using HiExM provides an opportunity to analyze nanoscale cellular features for a range of applications including cell phenotyping as well as drug and toxicology studies. As proof of concept, we analyzed the effects of the cardiotoxic chemotherapy agent doxorubicin (dox) on human induced pluripotent stem cell derived cardiomyocytes (hiPSC-CMs) cultured in 96 well plates (Fig. 3a).16 Commonly used dox concentrations in in vitro cardiotoxicity studies range between 0.1 μM to 10 μM, depending on experimental goals17–20. Although studies largely focus on electrophysiology, structural changes at the cellular level also provide critical information for assessing cardiotoxicity. For example, changes in nuclear morphology as well as DNA damage are associated with cellular stress and apoptosis, which are commonly observed at high dox concentrations (>200 nM) in in vitro studies using conventional microscopy.17–21 In contrast, heart tissue is potentially exposed to lower dox concentrations in clinical treatments.16,22,23 Because HiExM has the sensitivity to detect nanoscale cellular features, we compared nuclei using Hoechst signal intensity in CMs treated with 1 nM, 10 nM, 100 nM, and 1 µM dox after 24 hours compared with DMSO (Fig. 3b). Although we did not observe a significant change in overall nuclear morphology (Fig. S10), images showed increasing Hoechst density at the nuclear periphery in a dox concentration dependent manner.

Revised version: The ability to image cells at increased resolution using HiExM provides an opportunity to analyze nanoscale cellular features for a range of applications including cell phenotyping as well as drug and toxicology studies. As proof of concept, we analyzed the effects of the cardiotoxic chemotherapy agent doxorubicin (dox) on human induced pluripotent stem cell derived cardiomyocytes (hiPSC-CMs) cultured in 96 well plates (Fig. 3a).16 Commonly used dox concentrations in in vitro cardiotoxicity studies range between 0.1 μM to 10 μM, depending on experimental goals17–20. Although studies largely focus on electrophysiology, structural changes at the cellular level also provide critical information for assessing cardiotoxicity. For example, changes in nuclear morphology as well as DNA damage are associated with cellular stress and apoptosis, which are commonly observed at high dox concentrations (>200 nM) in in vitro studies using conventional microscopy.17–21 In contrast, heart tissue is potentially exposed to lower dox concentrations in clinical treatments.16,22,23 Because HiExM has the sensitivity to detect nanoscale cellular features, we compared nuclei using Hoechst signal intensity in CMs treated with 1 nM, 10 nM, 100 nM, and 1 µM dox after 24 hours compared with DMSO (Fig. 3b). Although we did not observe a significant change in overall nuclear morphology (Fig. S11), images showed increasing Hoechst density at the nuclear periphery in a dox concentration dependent manner.

HiExM shows altered nuclear morphology of cardiomyocytes treated with low doses of doxorubicin

a. hiPSC-CMs imaged post-expansion in HiExM. b. Example images of hiPSC-CM nuclei after dox treatment for 24 hrs before (top) and after (bottom) expansion. c. Example images of hiPSC-CM nuclei taken before (top) and after (bottom) expansion. d. Heatmaps of nuclei in c showing pixel groupings for subsequent analysis. Heatmaps were generated from the image mask after dilation, such that the outer edge represents a contour just beyond the nuclear periphery. e. Relative Hoechst intensity plotted as a function of pixel position relative to the edge of the dilated mask in pre-(top) and post-(bottom) expansion nuclei. Colored curves represent dox concentrations or DMSO (ctrl). Shaded regions represent SEM for n = 56, 71, 64, 92, and 62 nuclei in the pre-expansion images and n = 4 replicates in the expanded case (118, 111, 110, 113, and 77 nuclei analyzed for DMSO control, 1 nM, 10 nM, 100 nM, and 1 µM Dox, respectively). f. Insets of plots in e highlighting the nuclear periphery. g. Rate-of-change analysis of curves for pre-(left) and post-(right) expansion images. Outer edge values were obtained by determining the maximum values of the derivatives for each condition, and inner edge values were obtained by determining the minimum values of the derivatives within the domain 0 - 0.4. h. Curvature analysis for pre-(left) and post-(right) expansion data. Values represent the minima of the second derivative for curves within the domain 0 – 0.4. i. Average Hoechst intensities plotted as gradients for each dox concentration. Scale bar represents 10 percent distance from the edge of the nucleus to the center. Error bars represent SEM as in e. * denotes significance from an independent two-sample t-test (p < 0.05). *Scale bar represents biological scale assuming 4x expansion.

Earlier version: To quantify this effect, we measured the gradient slope of Hoechst intensity around the nuclear periphery. To this end, we masked 529 nuclei from background in dox treated and control samples (Fig. 3c). We selected a single optical slice that intersected the nucleus perpendicularly along the z-axis for our analysis to avoid out-of-focus signal at the nuclear edge (Fig. S11). To account for irregular nuclear shapes, we iteratively grouped pixels into rings for each nucleus starting with the edge of the mask (Fig 3d). The mask is dilated outward from the edge of the nucleus to include a background reading as reference. The average pixel intensity value for each ring represents the average Hoechst intensity as a function of relative distance from the edge of the nucleus to its center (Fig. 3e, Movie S3). Notably, HiExM revealed a peak of Hoechst intensity at the edge of the nucleus that is not resolved in pre-expansion images (Fig. 3f). A line-scan-based analysis showed the same trend (Fig. S12). To further evaluate these differences, we next calculated the first derivative of each curve and found the maximum and minimum values, which correspond to the gradient slopes on either side of the peak (Fig. 3g, Fig. S13). In this case, a higher magnitude of gradient slope reflects a ‘sharper’ transition from the nuclear edge to the background (outside) or from the nuclear edge to the nucleoplasm (inside). Post-expansion, the gradient slopes on both sides of the peak showed a significant increase at dox concentrations as low as 10 nM compared to DMSO, whereas no trend is observed in pre-expansion images (Fig. 3g).

Revised version: To quantify this effect, we measured the gradient slope of Hoechst intensity around the nuclear periphery. To this end, we masked 529 nuclei from background in dox treated and control samples (Fig. 3c). We selected a single optical slice that intersected the nucleus perpendicularly along the z-axis for our analysis to avoid out-of-focus signal at the nuclear edge (Fig. S12). To account for irregular nuclear shapes, we iteratively grouped pixels into rings for each nucleus starting with the edge of the mask (Fig 3d). The mask is dilated outward from the edge of the nucleus to include a background reading as reference. The average pixel intensity value for each ring represents the average Hoechst intensity as a function of relative distance from the edge of the nucleus to its center (Fig. 3e, Movie S3). Notably, HiExM revealed a peak of Hoechst intensity at the edge of the nucleus that is not resolved in pre-expansion images (Fig. 3f). A line-scan-based analysis showed the same trend (Fig. S13). To further evaluate these differences, we next calculated the first derivative of each curve and found the maximum and minimum values, which correspond to the gradient slopes on either side of the peak (Fig. 3g, Fig. S14). In this case, a higher magnitude of gradient slope reflects a ‘sharper’ transition from the nuclear edge to the background (outside) or from the nuclear edge to the nucleoplasm (inside). Post-expansion, the gradient slopes on both sides of the peak showed a significant increase at dox concentrations as low as 10 nM compared to DMSO, whereas no trend is observed in pre-expansion images (Fig. 3g).

Earlier version: In addition to gradient slopes at the nuclear edge, we also quantified the pattern of Hoechst signal at the nuclear edge by calculating the curvature of the peak using the second derivative (Fig. 3h). These data also showed a significant increase in Hoechst intensity at 10 nM dox compared to DMSO, whereas no significant difference is observed using conventional confocal microscopy (Fig. 3i). We note that, for both the outer gradient slope and peak curvature, values decrease from 100 nM to 1 µM dox, consistent with increased chromatin condensation associated with apoptosis typically observed at higher dox concentrations19,24. In contrast, CMs treated with increasing concentrations of H2O2, another cellular stressor, did not show similar differences in Hoechst patterns (Fig. S14). Although direct evidence specifically linking dox to increased DNA condensation at the nuclear periphery is limited, the known pro-apoptotic effects of dox strongly suggest that our observations correlate with these changes.17,19,25 Thus, we suggest that HiExM can detect early alterations in response to dox treatment using Hoechst signal, setting the stage for further investigation to understand the consequences across a range of drugs and conditions. Overall, our analysis demonstrates the ability of HiExM to detect subtle cellular changes compared to standard confocal microscopy that could improve the resolution of drug screening approaches.

Revised version: In addition to gradient slopes at the nuclear edge, we also quantified the pattern of Hoechst signal at the nuclear edge by calculating the curvature of the peak using the second derivative (Fig. 3h). These data also showed a significant increase in Hoechst intensity at 10 nM dox compared to DMSO, whereas no significant difference is observed using conventional confocal microscopy (Fig. 3i). We note that, for both the outer gradient slope and peak curvature, values decrease from 100 nM to 1 µM dox, consistent with increased chromatin condensation associated with apoptosis typically observed at higher dox concentrations19,24. In contrast, CMs treated with increasing concentrations of H2O2, another cellular stressor, did not show similar differences in Hoechst patterns (Fig. S15). Although direct evidence specifically linking dox to increased DNA condensation at the nuclear periphery is limited, the known pro-apoptotic effects of dox strongly suggest that our observations correlate with these changes.17,19,25 Thus, we suggest that HiExM can detect early alterations in response to dox treatment using Hoechst signal, setting the stage for further investigation to understand the consequences across a range of drugs and conditions. Overall, our analysis demonstrates the ability of HiExM to detect subtle cellular changes compared to standard confocal microscopy that could improve the resolution of drug screening approaches.

Earlier version: Our work demonstrates that HiExM is a simple and inexpensive method that can increase the throughput of expansion microscopy. Another recent method also reports a scalable ExM-based method with comparable throughput, albeit this approach requires the use of a custom cell culture microplate and manual imaging.8 Our device is fabricated with CNC milling, a common fabrication technique available in most academic settings, making our method widely accessible. The HiExM device can also be produced by injection molding which can further extend its accessibility and distribution in research settings. We have also adapted our device to fit cell culture plates of varying dimensions including a 24 well plate which can accommodate greater expansion and significantly higher resolution (Fig. S15).26 Moreover, the use of a photo-active gel solution enabled consistent expansion in our HiExM workflow. In this work, we found that the use of photoinitiation with irgacure caused signal loss for to varying degrees depending on the fluorophore; coupled with the signal dilution inherent to the expansion process, signal intensity is an important consideration when using HiExM. Our platform is amenable to other photo-active expansion chemistries such as the Photo-ExM chemistry which does not require CF dyes or an anoxic environment (Fig. S16)12,13.

Revised version: Our work demonstrates that HiExM is a simple and inexpensive method that can increase the throughput of expansion microscopy. Another recent method also reports a scalable ExM-based method with comparable throughput, albeit this approach requires the use of a custom cell culture microplate and manual imaging.8 Our device is fabricated with CNC milling, a common fabrication technique available in most academic settings, making our method widely accessible. The HiExM device can also be produced by injection molding which can further extend its accessibility and distribution in research settings. We have also adapted our device to fit cell culture plates of varying dimensions including a 24 well plate which can accommodate greater expansion and significantly higher resolution (Fig. S16).26 Moreover, the use of a photo-active gel solution enabled consistent expansion in our HiExM workflow. In this work, we found that the use of photoinitiation with irgacure caused signal loss for to varying degrees depending on the fluorophore; coupled with the signal dilution inherent to the expansion process, signal intensity is an important consideration when using HiExM. Our platform is amenable to other photo-active expansion chemistries such as the Photo-ExM chemistry which does not require CF dyes or an anoxic environment (Fig. S17)12,13.

Another advantage of HiExM is that it is fully compatible with autonomous imaging, allowing data acquisition of thousands of cells within a day. HiExM achieved a resolution of ∼115 nm in a 96-well culture plate using a 63x water immersion objective with a 1.15 numerical aperture. By employing microscopes with a higher numerical aperture, we expect our platform can achieve significantly greater resolution. We also show that HiExM error measurements are similar to proExM and suggest that expansion distortions should be considered for robust data interpretation. Recent innovative work that developed a microplate-based ExM protocol maintains hydrogel geometry and orientation improves coordinate-based imaging to accurately assess distortions with ease.27 Another recent work used micropatterning to apply expandable protein grids to coverslips for local expansion factor measurement. This method could be adapted to HiExM to enable expansion measurements in large datasets where non-rigid registration is impractical.28 Taken together, HiExM is an accessible nanoscale imaging platform designed for scalable sample preparation and autonomous imaging, broadening the use of ExM in biological research and pharmaceutical applications.

Added in the revised version: Supporting information